# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

```python
x='123'
y='321'
if sorted(y)==sorted(x):
    print('yes')
```

Output:
```
yes
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    if sorted(x)==sorted(y):
        return True
    else:
        return False
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
lst=[]
for i in range(1,1000000):
    f=str(i)
    g=str(int(f)*2)
    k=str(int(f)*3)
    h=str(int(f)*4)
    m=str(int(f)*5)
    c=str(int(f)*6)
    if same_digits(f,g)==True:
        lst.append(f)
    elif same_digits(f,k)==True:
        lst.append(f)
    elif same_digits(f,h)==True:
        lst.append(f)
    elif same_digits(f,m)==True:
        lst.append(f)
    elif same_digits(f,c)==True:
        lst.append(f)
min(lst)
```

Output:
```
'100035'
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```